<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/katas/algorithms/LinearRegression_BostonHousing.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting exercise generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

---

# Kata: Predict Housing Prices with Linear Regression

The goal of this kata is to predict the median price (in $1,000's) of homes given their characteristics.

The [Boston Housing Prices](http://www.cs.toronto.edu/~delve/data/boston/bostonDetail.html) dataset is frequently used to test regression algorithms.

![Boston suburb](images/boston_housing.jpg)

The dataset contains information gathered in the 1970s concerning housing in the Boston suburban area. Each house has the following features.

| Feature | Description |
| --- | --- |
| 0 | Per capita crime rate by town |
| 1 | Proportion of residential land zoned for lots over 25,000 sq.ft. |
| 2 | Proportion of non-retail business acres per town. |
| 3 | Charles River dummy variable (1 if tract bounds river; 0 otherwise) |
| 4 | Nitric oxides concentration (parts per 10 million) |
| 5 | Average number of rooms per dwelling |
| 6 | Proportion of owner-occupied units built prior to 1940 |
| 7 | Weighted distances to five Boston employment centres |
| 8 | Index of accessibility to radial highways |
| 9 | Full-value property-tax rate per $10,000 |
| 10 | Pupil-teacher ratio by town |
| 11 | 1000(Bk - 0.63)^2 where Bk is the proportion of blacks by town |
| 12 | Lower status of the population |



## Package setup

```python
# Import base packages
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
```

```python
# Setup plots
%matplotlib inline
plt.rcParams['figure.figsize'] = 10, 8
%config InlineBackend.figure_format = 'retina'
sns.set()
```

### Question

Import the needed packages.

<!-- grade_id: cell-540c08618645db39, points: 0, grade: false, solution: true -->
```python
# Import ML packages (edit this list if needed)
from sklearn.datasets import load_boston

from sklearn.model_selection import train_test_split
```

## Step 1: Loading the data

### Question

Load the Boston Housing dataset included with scikit-learn into the `dataset` variable.

<!-- grade_id: cell-9cd3e09cd8208cc9, points: 0, grade: false, solution: true -->
```python
dataset = load_boston()
```

```python
# Describe the dataset
print(dataset.DESCR)
```

Output:
```
.. _boston_dataset:

Boston house prices dataset
---------------------------

**Data Set Characteristics:**  

    :Number of Instances: 506 

    :Number of Attributes: 13 numeric/categorical predictive. Median Value (attribute 14) is usually the target.

    :Attribute Information (in order):
        - CRIM     per capita crime rate by town
        - ZN       proportion of residential land zoned for lots over 25,000 sq.ft.
        - INDUS    proportion of non-retail business acres per town
        - CHAS     Charles River dummy variable (= 1 if tract bounds river; 0 otherwise)
        - NOX      nitric oxides concentration (parts per 10 million)
        - RM       average number of rooms per dwelling
        - AGE      proportion of owner-occupied units built prior to 1940
        - DIS      weighted distances to five Boston employment centres
        - RAD      index of accessibility to radial highways
        - TAX      full-value property-tax rate per $10,000
        - PTRATIO  pu
```

```python
# Show a sample of raw training data
df_boston = pd.DataFrame(dataset.data)
# Add target to DataFrame
df_boston.columns = dataset.feature_names
df_boston['MEDV'] = dataset.target
df_boston.sample(n=10)
```

Output:
```
CRIM    ZN  INDUS  CHAS    NOX     RM    AGE     DIS   RAD    TAX  \
477  15.02340   0.0  18.10   0.0  0.614  5.304   97.3  2.1007  24.0  666.0   
23    0.98843   0.0   8.14   0.0  0.538  5.813  100.0  4.0952   4.0  307.0   
211   0.37578   0.0  10.59   1.0  0.489  5.404   88.6  3.6650   4.0  277.0   
81    0.04462  25.0   4.86   0.0  0.426  6.619   70.4  5.4007   4.0  281.0   
186   0.05602   0.0   2.46   0.0  0.488  7.831   53.6  3.1992   3.0  193.0   
390   6.96215   0.0  18.10   0.0  0.700  5.713   97.0  1.9265  24.0  666.0   
295   0.12932   0.0  13.92   0.0  0.437  6.678   31.1  5.9604   4.0  289.0   
35    0.06417   0.0   5.96   0.0  0.499  5.933   68.2  3.3603   5.0  279.0   
164   2.24236   0.0  19.58   0.0  0.605  5.854   91.8  2.4220   5.0  403.0   
303   0.10000  34.0   6.09   0.0  0.433  6.982   17.7  5.4917   7.0  329.0   

     PTRATIO       B  LSTAT  MEDV  
477     20.2  349.48  24.91  12.0  
23      21.0  394.54  19.88  14.5  
211     18.6  395.24  23.98  19.3  
81    
```

## Step 2: Preparing the data

### Question

Store input data and labels into the `x` and `y` variables respectively.

<!-- grade_id: cell-18c91823218fb4cd, points: 0, grade: false, solution: true -->
```python
x = dataset.data
y = dataset.target
```

<!-- grade_id: cell-e0357659e2fe5f21, points: 1, grade: true, solution: false -->
```python
print(f'x: {x.shape}. y: {y.shape}')

assert x.shape == (506, 13)
assert y.shape == (506,)
```

Output:
```
x: (506, 13). y: (506,)
```

### Question

Prepare data for training. Store the data subsets in variables named `x_train`/`y_train` and `x_test`/`y_test` with a 20% ratio.

<!-- grade_id: cell-0110545e002f2425, points: 0, grade: false, solution: true -->
```python
x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=.2)
```

<!-- grade_id: cell-6bdb09815ea10759, points: 1, grade: true, solution: false -->
```python
print(f'x_train: {x_train.shape}. y_train: {y_train.shape}')
print(f'x_test: {x_test.shape}. y_test: {y_test.shape}')

assert x_train.shape == (404, 13)
assert y_train.shape == (404,)
assert x_test.shape == (102, 13)
assert y_test.shape == (102,)
```

Output:
```
x_train: (404, 13). y_train: (404,)
x_test: (102, 13). y_test: (102,)
```

### Question

Scale features by standardization while preventing information leakage from the test set.

<!-- grade_id: cell-1f765aa560e5a2d3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-ee2cab814e063754, points: 1, grade: true, solution: false -->
```python
mean_train = x_train.mean()
std_train = x_train.std()
print(f'mean_train: {mean_train}. std_train: {std_train}')

assert np.abs(np.max(mean_train)) < 10**-6
assert np.abs(np.max(std_train - 1)) < 10**-6
```

## Step 3: Training a model

### Question

Create a `SGDRegressor` instance and store it into the `model` variable. Fit this model on the training data.

<!-- grade_id: cell-7a309bc803a9f6e2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

## Step 4: Evaluating the model

### Question

Compute the training and test MSE into the `mse_train` and `mse_test` variables respectively.

<!-- grade_id: cell-6b0051e9c4adbc2b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-38a4956fe588ef5a, points: 1, grade: true, solution: false -->
```python
print (f'Training MSE: {mse_train:.2f}. Test MSE: {mse_test:.2f}')

plt.scatter(y_test, y_test_pred)
plt.xlabel("Actual Prices")
plt.ylabel("Predicted Prices")
plt.title("Actual Prices vs Predicted Prices")
```

### Question

Go back to step 3 and try to obtain the best possible test MSE by tweaking the `SGDRegressor` parameters.

## Step 5: Use another regression algorithm

### Question

- Create and fit a `LinearRegression` instance, which uses the normal equation instead of gradient descent.
- Compute the training and test MSE for this instance (variables `mse_train_n` and `mse_test_n`). How does it compare to the `SGDRegressor` in this case?

<!-- grade_id: cell-e3e4339cefd72706, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-559b88ad41545a83, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3bccb3c8d7805572, points: 1, grade: true, solution: false -->
```python
print (f'Training MSE: {mse_train_n:.2f}. Test MSE: {mse_test_n:.2f}')

plt.scatter(y_test, y_test_pred_n)
plt.xlabel("Actual Prices")
plt.ylabel("Predicted prices")
plt.title("Actual Prices vs Predicted Prices")
```

```python

```